---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Case Study of Investment Portfolio Risk Management and Optimization

```yaml
type: VideoExercise
key: 8e69e04291
xp: 50
```

`@projector_key`
6cbd71baf5e0b248b2939f892098f9e9
